evaluation_interval: 1m

rule_files:
- caching_availability_recording_rules.yaml

tests:
  - interval: 1m
    # All replicas up → konflux_up = 1
    input_series:
      - series: 'squid_up{namespace="proxy", service="squid", job="uwm/scrape", instance="10.0.0.1:9301", pod="squid-0"}'
        values: '1 1 1'
      - series: 'squid_up{namespace="proxy", service="squid", job="uwm/scrape", instance="10.0.0.2:9301", pod="squid-1"}'
        values: '1 1 1'
    promql_expr_test:
      - expr: konflux_up
        eval_time: 2m
        exp_samples:
          - labels: 'konflux_up'
            value: 1

  - interval: 1m
    # One replica down → konflux_up = 0 (min across targets)
    input_series:
      - series: 'squid_up{namespace="proxy", service="squid", job="uwm/scrape", instance="10.0.0.1:9301", pod="squid-0"}'
        values: '1 1 1'
      - series: 'squid_up{namespace="proxy", service="squid", job="uwm/scrape", instance="10.0.0.2:9301", pod="squid-1"}'
        values: '1 0 0'
    promql_expr_test:
      - expr: konflux_up
        eval_time: 2m
        exp_samples:
          - labels: 'konflux_up'
            value: 0
